# Laboratorio 1

Bienvenido al primer laboratorio del curso Análisis de Datos en Python. 

Escribe el código correspondiente para contestar cada una de las preguntas. Fíjate muy bien en el nombre de las variables, ya que de lo contrario, tu trabajo no podrá ser evaluado de manera correcta. 

Debes ir ejecutando las celdas de codigo en el orden que aparecen, para evitar tener errores. 

Este trabajo tiene un límite de tiempo que puedes visualizar al inicio de esta pantalla. 

```python
# Primero importamos las librerias y los datos que utilizaremos.

import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
import seaborn as sns
import datetime
import matplotlib.pyplot as plt
import warnings
warnings.filterwarnings('ignore')
df = pd.read_csv('data_football_ratings.csv')
```

# Pregunta 1
¿Cuál es el promedio, mediana y desviación estándar de los datos registrados en la columna “original_rating”? 

Deben crear el código necesario para el promedio, la mediana y la desviación standard
*almacena las respuestas en formato int en las variables promedio, mediana y desvest respectivamente*

Usa la funcion `int()` para convertir a entero. 

```python
df.head()
```

Output:
```
competition        date                    match     team  pos pos_role  \
0   Euro 2016  10/06/2016  France - Romania, 2 - 1  Romania   DF       DC   
1   Euro 2016  10/06/2016  France - Romania, 2 - 1  Romania   DF       DC   
2   Euro 2016  10/06/2016  France - Romania, 2 - 1  Romania   DF       DC   
3   Euro 2016  10/06/2016  France - Romania, 2 - 1   France  Sub      Sub   
4   Euro 2016  10/06/2016  France - Romania, 2 - 1  Romania   MF      DMC   

            player      rater  is_human  original_rating  ...  \
0   Dragos Grigore     Kicker         1             3.50  ...   
1   Dragos Grigore  WhoScored         0             6.56  ...   
2   Dragos Grigore  SofaScore         0             6.70  ...   
3  Anthony Martial  WhoScored         0             6.19  ...   
4   Mihai Pintilii     Kicker         1             3.50  ...   

   betweenness_centrality  closeness_centrality  flow_centrality  \
0                0.143055              0.603571         0.304348   
1           
```

### Pregunta 2
¿Cuántos de los partidos de la base de datos tuvieron una duración (“game_duration”) de 120 minutos?
*Alamacena la respuesta en formato int en la variable rep_P2*

<!-- grade_id: cell-a9e742e4c226ea1a, points: 0, grade: false, solution: true -->
```python
rep_P2=0
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-a2495efa7628513a, points: 10, grade: true, solution: false -->
```python

```

### Pregunta 3
Considerando solamente los partidos de la Copa del Mundo 2018 (“World Cup 2018”), ¿cuántos partidos disputó Croacia?
*almacena tu respuesta en la variable n_croacia*

<!-- grade_id: cell-c4ec5d3f2b70991e, points: 0, grade: false, solution: true -->
```python
n_croacia = 0
# your code here
raise NotImplementedError
```

<!-- grade_id: cell-fcdc878340735073, points: 10, grade: true, solution: false -->
```python

```

### Pregunta 4
Construya un gráfico de torta con los goles marcados por Robert Lewandowski, agrupando los goles que marcó por su selección y por su club.

**Importante considerar que para esta pregunta se revisarán manualmente las entregas der manera aleatoria, quién responda sin realizar el gráfico solicitado tendrá 0 puntos.**

<!-- grade_id: cell-120b7b8fcfc97fd1, points: 0, grade: false, solution: true -->
```python
## Usen esta celda para crear el grafico. 

# your code here
raise NotImplementedError
```

### A partir del gráfico construido, complete la información de las siguientes variables en formato int: "Goles_Bayern","Goles_Poland","Porcentaje_Bayern","Porcentaje_Poland"

Los porcentajes se aproximan al entero más cercano, por ejemplo 97.32% será 0.97 o 7.85% será 0.08

<!-- grade_id: cell-ab0b8fac7d11b0ef, points: 0, grade: false, solution: true -->
```python

Goles_Bayern=0
Goles_Poland=0
Porcentaje_Bayern=0
Porcentaje_Poland=0


# your code here
raise NotImplementedError
```

<!-- grade_id: cell-70762188e9a9150e, points: 10, grade: true, solution: false -->
```python

```

### Pregunta 5
¿Cuántas tarjetas rojas fueron mostradas en la Eurocopa 2016 (“Euro 2016”)? *Almacena tu respuesta en la variable rcards*

<!-- grade_id: cell-ca547a367e621187, points: 0, grade: false, solution: true -->
```python
rcards = 0

# your code here
raise NotImplementedError
```

<!-- grade_id: cell-cf57b3130ebb2e08, points: 10, grade: true, solution: false -->
```python

```

### Pregunta 6
Construya un Dataframe que contenga, para cada jugador del archivo, la cantidad total de partidos en los que fue evaluado. Muestre solamente los 10 primeros registros (según orden alfabético del nombre del jugador) de este DataFrame. Considere que un jugador no es evaluado cuando original_rating = 1 y los minutos jugados por el jugador es menor que 20 (minutos).

Indicaciones: 
- el nombre de las columnas del dataframe no importa
- el dataframe solo debe tener 2 columnas, la primera con el nombre del jugador y la segunda con la cantidad de partidos en los que fue evaluado
- almacena el dataframe en la variable df_p6
- ordena el dataframe según el nombre de los jugadores, de manera ascendente (en general al hacer el group by deberia quedar con este orden!)
- *Hint: los métodos .head() y .groupby() puede ser útiles*. Cuidado con los duplicados!!

<!-- grade_id: cell-1d89c7054a6e49f8, points: 0, grade: false, solution: true -->
```python

df_p6= df

# your code here
raise NotImplementedError
```

Probemos el resultado aqui: 

```python
df_p6.head(1)
```

La celda anterior debe tener los datos: 
**Aaron Cresswell 34**

**Cuidado!! Deja a los jugadores como indice!!**

No crees un indice numerico nuevo. En caso de que se haya creado automaticamente, usa el comando

``` df_p6.set_index('player', inplace=True)```

donde `player` corresponde al nombre de la variable donde guardas los nombres de los jugadores

A continuación, hay un test visible, para que puedas ver que tu dataframe este bien creado. Este test vale 1 punto.

<!-- grade_id: cell-3037a1d13268ba68, points: 1, grade: true, solution: false -->
```python
assert df_p6.head(1).values == 34, "Revisa el formato de tu DataFrame. El primer registro debe ser Aaron Cresswell y los partidos en que fue evaluado son 34"
```

La última celda corresponde a un test oculto, que revisa el dataframe pedido en la pregunta 6. **El dataframe solo debe contener los primeros 10 registros. De lo contrario, no tendrás puntaje.**  

```python
## Revisa que tu dataframe tenga el largo correcto!! 

len(df_p6) == 10
```

<!-- grade_id: cell-48e9c954bcf215da, points: 9, grade: true, solution: false -->
```python


```